# Question 3

In this question, we'll focus on looping.

### Part A

In the question below, you'll write code that generates a new list, `reverse`, that is the same as the input list except with its elements backwards relative to the original.

As an example, if the input list is `[1, 2, 3]`, the output list `reverse` should be `[3, 2, 1]`.

**You have to use loops** to answer this question! `[::-1]` is a neat trick that you should definitely use, and `reversed` is a great function you should also definitely use, but you just can't use them here.

<!-- grade_id: q3a, points: 0, grade: false, solution: true -->
```python
def reverse_list(inlist):
    reverse = []

    ### BEGIN SOLUTION
    for i in inlist:
        reverse=inlist[-i]
    ### END SOLUTION
    return reverse
```

Output:
```
<ipython-input-28-bd5a5fbc996e> in <module>
      8     return reverse
      9 oops=[1,2,3]
---> 10 reverse_list(oops)

<ipython-input-28-bd5a5fbc996e> in reverse_list(inlist)
      3 
      4     ### BEGIN SOLUTION
----> 5     for i in len(inlist):
      6         reverse=inlist[-i]
      7     ### END SOLUTION

TypeError: 'int' object is not iterable
```

<!-- grade_id: q3a_test1, points: 7, grade: true, solution: false -->
```python
t1 = [1, 2, 3]
t2 = [3, 2, 1]
assert t2 == reverse_list(t1)
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-26-e7a5cfeb46a4> in <module>
      1 t1 = [1, 2, 3]
      2 t2 = [3, 2, 1]
----> 3 assert t2 == reverse_list(t1)

AssertionError:
```

<!-- grade_id: q3a_test2, points: 7, grade: true, solution: false -->
```python
t1 = [2, 6, 52, 64, 3.4, "something", 3, 6.2]
t2 = [6.2, 3, "something", 3.4, 64, 52, 6, 2]
assert t1 == reverse_list(t2)
```

### Part B

This time, you'll implement a method for computing the average value of a list of numbers. Remember how to compute an average: add up all the values, then divide that sum by the number of values you initially added together.

Example Input: [2, 1, 4, 3]

Example Output: 2.5

Aside from `len()`, you **cannot** use built-in functions like `sum()` or import any packages to help you out. You have to write this on your own!

<!-- grade_id: q3b, points: 0, grade: false, solution: true -->
```python
def average(numbers):
    avg_val = 0.0
    
    ### BEGIN SOLUTION
    summa=0.0
    for i in numbers:
        summa+=i
        avg_val=summa/len(numbers)
    
    ### END SOLUTION
    
    return avg_val
```

<!-- grade_id: q3b_test1, points: 7, grade: true, solution: false -->
```python
import numpy as np

inlist = np.random.randint(10, 100, 10).tolist()
np.testing.assert_allclose(average(inlist), np.mean(inlist))
```

<!-- grade_id: q3b_test2, points: 7, grade: true, solution: false -->
```python
inlist = np.random.randint(10, 1000, 10).tolist()
np.testing.assert_allclose(average(inlist), np.mean(inlist))
```

### Part C

In this question, the function below `first_negative` takes a list of numbers (creatively-named `numbers`). Your goal is to write a `while` loop that searches for the **first negative number** in the list.

You should set the variable `num` equal to that negative number.

If there is no negative number in the list (all the numbers are `>= 0`), just return 0.

<!-- grade_id: q3c, points: 0, grade: false, solution: true -->
```python
def first_negative(numbers):
    num = 0
    ### BEGIN SOLUTION
    count = 0
    while count < len(numbers):
        if numbers[count]<0:
            return numbers[count]
        count = count+1
    ### END SOLUTION
    
    return num
```

<!-- grade_id: q3c_test1, points: 5, grade: true, solution: false -->
```python
inlist1 = [1, 4, -22, -39, 1, 8, 27, 6, 27, -12]
a1 = -22
assert a1 == first_negative(inlist1)
```

<!-- grade_id: q3c_test2, points: 2, grade: true, solution: false -->
```python
inlist2 = [-29, -1, 49, -33, 49, -40, -7, 33, 36, 41]
a2 = -29
assert a2 == first_negative(inlist2)
```

```python

```